# MEEN 357 - Summer 2025
### Submission Instructions

- **Run All Cells**: Before submitting, go to **Kernel > Restart Kernel & Run All Cells** to ensure your code runs without errors. Submissions with errors will receive a **ZERO** grade.
- **Enter Your Name**: Fill in your name in the provided cell.
- **Complete the Code**: Replace all instances of `YOUR CODE HERE` with your solution. Remove `raise NotImplementedError()`.
- **Maintain Structure**: Do not add or remove any cells.
- **Test Your Code**: Run the provided tests to check your answers. Note that additional hidden tests may be used during grading.
- **Partial Credit**: Will be awarded only if your code runs error-free.
- **Save and Submit**: Ensure you submit the latest, correct version of your assignment by checking the last modified time.


```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3.8, "Your version of IPython is too old, please update it."
```

---

# Solving FEM

This assignment focuses on solving an FEM problem involving a discrete system of elements. The first part of the assignment is for you to get familiar with a new data type called pandas DataFrame. You will learn how to create and use a Pandas DataFrame. You will need to know this before proceeding with the assignment. 

The second part of the assignment breaks down the solving of a discrete element system into multiple steps. Every step is depended on the results from the previous step. Hence, make sure that you have the correct results before proceeding.

* [ Creating a pandas DataFrame](#Creating-a-pandas-DataFrame-(4-points))
* [ Pandas apply method](#Pandas-apply-method-(4-points))


### FEM for discrete spring systems
* [ Create the list of elements, nodes and boundary conditions](#Create-the-list-of-elements,-nodes-and-boundary-conditions-(6-points))
* [ Find element length and stiffness](#Find-element-length-and-stiffness-(5-points))
* [ Finding the global stiffness matrix](#Finding-the-global-stiffness-matrix-(6-points))
* [ Partition the global stiffness matrix](#Partition-the-global-stiffness-matrix-(3-points))
* [ Establish the displacement and load vectors](#Establish-the-displacement-and-load-vectors-(3-points))
* [ Solve for the unknowns](#Solve-for-the-unknowns-(7-points))
* [ Creating the Final function](#Create-an-FEM-function-(7-points))


```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

## Creating a pandas DataFrame (4 points)

Write a script that creates a pandas DataFrame called  **persons** with the following columns

* **name**, the name of the person
* **height**, the height of the person
* **weight**, the weight of the person
    

with the following 3 entries:

|index|name |height| weight|
|---|---|---|---|
|0| John| 160| 80|
|1| Alice| 150| 60|
|2| Bob| 155| 70|

**HINT:**
* Create a list of lists to contain the data. 

<!-- grade_id: dict_script, points: 0, grade: false, solution: true -->
```python
columns = ['columnOne','columnTwo']
data =    [[1, 10],
           [2, 20]]
df = pd.DataFrame(data, columns = columns)
# YOUR CODE HERE
raise NotImplementedError()
```

### Tests

<!-- grade_id: dict_test, points: 4, grade: true, solution: false -->
```python
columns = ['name', 'height', 'weight']
data =    [['John', 160, 80]]
persons_correct = pd.DataFrame(data, columns = columns)
assert ((persons_correct.iloc[0].sort_index()==persons.iloc[0].sort_index()).all()), ' Your persons list of dictionaries is wrong '
print('All correct. Good work !')
```

## Pandas apply method (4 points)

Write a script that can applies a function on each entry in the DataFrame called **persons** that you created in the previous problem.
Calculate the Body mMass Index, **BMI** for each person on the list using the following formula:

$BMI = \frac{weight}{height}$

Add a new column to the **persons** DataFrame called **BMI** to store the calculated values.

**HINT:**
* Use the following code to apply a function `compute(data)` on each row of a DataFrame **df** 

`result = df.apply(compute, axis = 1)`

* The result can then be assigned back into the DataFrame as a new column using the code:

`df['NewColumn'] = result`

<!-- grade_id: apply_script, points: 0, grade: false, solution: true -->
```python
def compute(data):
    pass
persons['computed_column'] = persons.apply(compute, axis=1)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: apply_test, points: 4, grade: true, solution: false -->
```python
columns = ['name', 'height', 'weight', 'BMI']
data =    [['John', 160, 80, 0.5]]
persons_correct = pd.DataFrame(data, columns = columns)
assert ((persons_correct.iloc[0].sort_index()==persons.iloc[0].sort_index()).all()), ' Your persons list of dictionaries is wrong '
print('All correct. Good work !')
```

# FEM for discrete systems

## Create the list of elements and nodes (6 points)

Create a program to solve the following FEM problem. Use the figure below as a strict guide for numbering your elements and nodes.

<img src="problem1.png" width="500"/>

To begin, write a script to complete the following:

1. Create a Pandas DataFrame called **nodes** with the following columns
    * **coordinate**, the x-coordinate of the node
    * **displacement**, the displacement for each node. Enter the **Displacement Boundary Conditions** here. For the unknown displacements, use `np.nan`.
    * **load**, the load for each node. Enter the **Load Boundary Conditions** here. For the unknown load, use `np.nan`.
    
1. Create a Pandas DataFrame called **elements** with the following columns
    * **start**, the starting node (use the left node of the element)
    * **end**, the end node (use the right node of the element)
    * **area**, the area of cross-section
    * **material**, the Modulus of elasticity
    
    The area and material of each element is given below. For the rest of the information, rely on the figure above.

| index|area|material|
|---|---|---|
| 0|4|14|
| 1|6|12|
| 2|6|15|

    
    Use the figure above to find the x-coordinates, and boundary conditions of the nodes:

<!-- grade_id: inputs_script, points: 0, grade: false, solution: true -->
```python
# Elements
columns = ['start', 'end'] #.... Complete these lines
data =    [[2, 0],
           ]
elements = pd.DataFrame(data, columns = columns)
# Nodes

# YOUR CODE HERE
raise NotImplementedError()
```

### Tests

<!-- grade_id: inputs_test, points: 6, grade: true, solution: false -->
```python
elements_correct = pd.Series({'start':0,'end':1,'area':4,'material':14})
assert ((elements.loc[0].sort_index() == elements_correct.sort_index())).all(), 'Your script is probably wrong'
print('All correct. Good work !')
```

## Find element length and stiffness (5 points)

Write a script that applies functions on each row of the Pandas DataFrame **elements** and computes the length and stiffness of the element. Use the formula shown below to calculate the length and the stiffness. Create new columns called **length** and **stiffness** in the same DataFrame to store the calculated values.

Length , $l=|x_j-x_i|$

Where $ x_i $ is the x-coordinate of the $i$ th node

$ i $ and $j$ are the index of the start and end node of the element, respectively


Stiffness, $k = \frac{EA}{l}$

Where
$E$ is the modulus of elasticity of the element, and 

$A$ is the cross sectional area of the element




<!-- grade_id: len_stiff_script, points: 0, grade: false, solution: true -->
```python
def computeLength(element):
    pass # Remove this as well

def computeStiffness(element):
    pass

elements['length'] = elements.apply(computeLength, axis=1)
# Modify/Complete the rest of the code.

# YOUR CODE HERE
raise NotImplementedError()
```

### Tests

<!-- grade_id: len_stiff_test, points: 5, grade: true, solution: false -->
```python
correct_stiff = 18
assert (abs(elements.iloc[2]['stiffness'] - correct_stiff) < 1e-4), ' Your script is probably wrong'
print('All correct. Good work !')
```


## Finding the global stiffness matrix (6 points)

Each element can be represented as follows
![](element.png)
Where $i$ and $j$ are the start and end node of element $e$,

$P_i$ is the force at node $i$,

$U_i$ is the displacement of node $i$,

$ K_e$ is the stiffness of the element $e$

The local set of equations can be written in matrix form as

$ \left[\matrix{ K_e & -K_e \\
                 -K_e & K_e } \right] \left( \matrix{U_i \\ U_j } \right) = \left( \matrix{P_i \\ P_j } \right)$ 

The global stiffness matrix is then formed by adding up the local stiffness matrix using the correct indices. For the example shown below, 

![](2element.png)

The global set of equations will be

$ \left[\matrix{ K_0 & -K_0 & 0 \\
                -K_0 & K_0+K_1 & -K_1 \\
                 0 & -K_1 & K_1 } \right] \left( \matrix{U_0 \\ U_1 \\ U_2 } \right) = \left( \matrix{P_0 \\ P_1 \\P_2 } \right)$ 

Write a script that takes information from the Pandas DataFrame called **elements** that you created in the previous task, and then generates the Global stiffness matrix, **K**.

**HINT:**

* You can create multiple global K matrices with zeros, and populate it with the correct values at the right indicies for each element. And then, sum up these matrices.

* You need to only change those parts of the global stiffness matrix that correspond to the index of the start and end nodes of the element. Therefore, use indexing. For example, if the start and end nodes are 0, 2, then you need to update positions, `(0,0), (0,2), (2,0), and (2,2)` in the K matrix.

`rows = [0,2]
columns = [0,2]
A[np.ix_(rows, columns)] = new_values`

<!-- grade_id: stiff_script, points: 0, grade: false, solution: true -->
```python
# Compute the global stiffness matrix K for each element
def computeGlobalStiffness(element):
    pass
globalK_each = elements.apply(computeGlobalStiffness,axis=1)
K = globalK_each.sum()

# YOUR CODE HERE
raise NotImplementedError()
```

### Tests

<!-- grade_id: stiff_test, points: 6, grade: true, solution: false -->
```python
correct_K = np.array([[18.4, -7.2],[-7.2, 25.2]])
assert (abs(K[1:3,1:3]-correct_K) < 1e-4).all(), ' Your script is probably wrong'
print('All correct. Good work !')
```

## Partition the global stiffness matrix (4 points)


![](partition.png)

The displacement and load vectors consists of known (boundary conditions) and unknown values. We can split the system of equations and solve for parts of the unknowns. Write a function called **partition_K** to split the K matrix you found earlier into the 4 matrices indicated in the figure above. The function takes the following inputs:
* **K** is the stiffness matrix
* **A** is a boolean vector indicating if the displacement is unknown
* **B** is a boolean vector indicating if the load is unknown

**HINT:**
* You can get boolean vector using the `pd.isna()` function. This function indicates which rows in a pandas dataframe contains empty cells. You can apply it directly to a column using `df['column'].isna()`

* Use indexing. For example, if you would like to get the 4 elements, `(0,0), (0,2), (2,0), and (2,2)` from a matrix M, then use the code:

`A = [True, False, True]
B = [True, False, True]
MAB = M[np.ix_(A, B)]`

<!-- grade_id: partition_script, points: 0, grade: false, solution: true -->
```python
def partition_K(K, A, B):
    pass
A = nodes['displacement'].isna()
B = 0 # modify this line
#... Modify/Complete the rest of the code 


# YOUR CODE HERE
raise NotImplementedError()
KAA, KAB, KBA, KBB = partition_K(K, A, B)
KAA
```

### Tests

<!-- grade_id: partition_test, points: 4, grade: true, solution: false -->
```python
correct_KAA = np.array([[18.4, -7.2],[-7.2, 25.2]])
A = [False, True, True, False]
B = [True, False, False, True]
KAA,_,_,_ = partition_K(K, A, B)
assert np.allclose(KAA, correct_KAA), ' Your script is probably wrong'
print('All correct. Good work !')
```

## Solve for the unknowns (7 points)

Create a script that solves for the unknown displacement and load. You can solve for them using the partitioned stiffness matrices, the known displacements, and known loads. 

* Set up equation 3 and 4 shown below to solve for the unknowns:

<img src="solve.png" width="500"/>


* Create a copy of the **nodes** DataFrame and call it **result**. Input the solved values into the table to replace the NaN cells. 
* Add a new column to the **result** DataFrame called **new_coordinate**. This is the new location of the node after deformation and can be found by adding the solved displacements to the original coordinate. 
* Compute the sum of all the load and verify that it is ~0.

**HINT:**
* Use the function `np.linalg.inv()` to find the inverse of a matrix.
* Use the funciton `np.dot(A,B)` to do a matrix multiplication of 2 arrays
* Use the function `df['column'].sum()` to compute the sum of a column in a DataFrame.


<!-- grade_id: solve_script, points: 0, grade: false, solution: true -->
```python
# Known displacement, UB
UB = nodes.loc[B,'displacement']
# Known load, PA
PA = 0 # Modify this line
result = nodes.copy()
# Place the solved values into the results DataFrame
# result.loc[A, 'displacement'] = ....

# Modify/Complete the rest of the code

# YOUR CODE HERE
raise NotImplementedError()
```

### Tests

<!-- grade_id: solve_test, points: 7, grade: true, solution: false -->
```python
U_correct = np.array([ 0.        ,  0.54195804, -0.003885  ,  0.        ])
P_correct = np.array([-6.06993007, 10.        , -4.        ,  0.06993007])
assert np.allclose(result['displacement'], U_correct), ' Your script is probably wrong. Wrong results for displacements'
assert np.allclose(result['load'], P_correct), ' Your script is probably wrong. Wrong results for forces'
print('All correct. Good work !')
```