# MEEN 357 - Summer 2025
### Submission Instructions

- **Run All Cells**: Before submitting, go to **Kernel > Restart Kernel & Run All Cells** to ensure your code runs without errors. Submissions with errors will receive a **ZERO** grade.
- **Enter Your Name**: Fill in your name in the provided cell.
- **Complete the Code**: Replace all instances of `YOUR CODE HERE` with your solution. Remove `raise NotImplementedError()`.
- **Maintain Structure**: Do not add or remove any cells.
- **Test Your Code**: Run the provided tests to check your answers. Note that additional hidden tests may be used during grading.
- **Partial Credit**: Will be awarded only if your code runs error-free.
- **Save and Submit**: Ensure you submit the latest, correct version of your assignment by checking the last modified time.


```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3.8, "Your version of IPython is too old, please update it."
```

---

# Root finding methods

### Define errors
* [Relative Error](#Relative-error) (5 points)
* [True Error](#True-error) (5 points)

### Bracketing methods
* [Bisection method](#Bisection-method) (10 points)
* [False position](#False-position-method) (10 points)

### Open methods
* [Newton Raphson method](#Newton-Raphson-method) (10 points)
* [Secant method](#Secant-method) (10 points)


```python
import numpy as np
import pandas as pd
```

## Define Error functions

## Relative error
Write a function called **relativeError** that takes the following inputs:

* **x**, the new value
* **xold**, the previous value

And returns the following output:

* **error**, the calculated relative error

The formula for the relative error is shown below:

$e_{relative} = |\frac{x_{new}-x_{old}}{x_{new}}|100 \%$

<!-- grade_id: relError, points: 0, grade: false, solution: true -->
```python
def relativeError(x,xold):
    """
    Calculate the relative error between a current value and a previous value.

    The relative error is computed as the absolute value of the difference between
    the current value and the previous value, divided by the current value, and then
    multiplied by 100 to express it as a percentage.

    Parameters:
    -----------
    x : float
        The current value.
    xold : float
        The previous value.

    Returns:
    --------
    float
        The relative error between `x` and `xold`, expressed as a percentage.

    Example:
    --------
    >>> relativeError(10, 9)
    10.0
    
    >>> relativeError(5, 4.5)
    10.0
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f2ecf4da621be48c, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: relError_test, points: 5, grade: true, solution: false -->
```python
QandAs = [[49, 20, 59.183673469387756],
[26, 3, 88.46153846153845],
[78, 87, 11.538461538461538],
[84, 94, 11.904761904761903],
[11, 42, 281.8181818181818]]
for xnew, xold, Answer in QandAs: 
    studentAnswer = relativeError(xnew, xold)
    assert np.isclose(studentAnswer, Answer), f' Did not work for the inputs {xnew=},{xold=}. Expected {Answer=}. Got {studentAnswer=}'

print("All correct. Good work!")
```

## True error
Write a function called **trueError** that takes the following inputs:

* **x**, the new value
* **xtrue**, the previous value

And returns the following output:

* **error**, the calculated relative error

The formula for the relative error is shown below:

$e_{true} = |\frac{x-x_{true}}{x_{true}}|100 \%$

<!-- grade_id: trueError, points: 0, grade: false, solution: true -->
```python
def trueError(x,xtrue):
    """
    Calculate the true error between an estimated value and the true value.

    The true error is computed as the absolute value of the difference between
    the estimated value and the true value, divided by the true value, and then
    multiplied by 100 to express it as a percentage.

    Parameters:
    -----------
    x : float
        The estimated value.
    xtrue : float
        The true or actual value.

    Returns:
    --------
    float
        The true error between `x` and `xtrue`, expressed as a percentage.

    Example:
    --------
    >>> trueError(9.5, 10)
    5.0
    
    >>> trueError(4.5, 5)
    10.0
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d6d0d2ae9fc49542, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: trueError_test, points: 5, grade: true, solution: false -->
```python
QandAs = [[38, 35, 8.571428571428571],
[72, 44, 63.63636363636363],
[58, 38, 52.63157894736842],
[90, 60, 50.0],
[0, 57, 100.0]]
for x, xold, Answer in QandAs: 
    studentAnswer = trueError(x, xold)
    assert np.isclose(studentAnswer, Answer), f' Did not work for the inputs {x=},{xold=}. Expected {Answer=}. Got {studentAnswer=}'
print("All correct. Good work!")
```

## Bracketing methods

## Bisection method
Write a function called “bisection” that takes the following inputs
* **f** the equation whose root is to be determined
* **xl** the lower bound of the search interval
* **xu** the upper bound of the search interval
* **e_tolerance** the error tolerance in percentages
* **max_iteration** the maximum numbers of iterations

It then returns a Pandas DataFrame called **result** which contains the results from all iterations and has the following columns.
* **iteration** the iteration number
* **xl** the lower bound of the search interval
* **xu** the upper bound of the search interval
* **xr** the new root
* **relative error** the calculated relative error

Your function needs to compute the root iteratively using the bisection method until the error tolerance condition has been met or the number of iterations has exceeded the maximum number of iterations. 

Using the bisection method, the new root, $x_r$ is computed as $x_r = \frac{x_l+x_u}{2}$

<!-- grade_id: bisection, points: 0, grade: false, solution: true -->
```python
def bisection(f, xl, xu, e_tolerance, max_iterations): 
    """
    Perform the bisection method to find the root of a function within a given interval.

    The bisection method iteratively narrows the interval [xl, xu] where the root of the function f(x) is located.
    It stops when the relative error falls below a specified tolerance or the maximum number of iterations is reached.

    Parameters:
    -----------
    f : function
        The function for which the root is to be found.
    xl : float
        The lower bound of the interval.
    xu : float
        The upper bound of the interval.
    e_tolerance : float
        The stopping criterion for the relative error (as a percentage).
    max_iterations : int
        The maximum number of iterations to perform.

    Returns:
    --------
    pandas.DataFrame
        A DataFrame containing the details of each iteration, including:
        - 'xl': The current lower bound of the interval.
        - 'xu': The current upper bound of the interval.
        - 'xr': The midpoint of the current interval (the approximate root).
        - 'relative error': The relative error at the current iteration (as a percentage).

    Example:
    --------
    >>> def func(x):
    >>>     return x**2 - 4

    >>> bisection(func, 1, 3, 0.01, 100)
          xl   xu   xr  relative error
    1     1.0  3.0  2.0        50.000000
    2     1.0  2.0  1.5        33.333333
    ...

    Notes:
    ------
    - The function assumes that f(xl) and f(xu) have opposite signs, indicating that a root lies within the interval [xl, xu].
    - The loop will terminate early if the function converges before reaching the maximum number of iterations.
    - The function uses the `relativeError` function to calculate the relative error.

    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ec3e7957b793484d, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: bisection_test, points: 10, grade: true, solution: false -->
```python
func = lambda x: 667.38/x*(1-np.exp(-0.146843*x))-40;
QandAs = [[3, 18, 0.3, 6, 14.953125],
[7, 17, 0.2, 8, 14.7734375],
[10, 23, 0.4, 5, 14.46875],
[10, 21, 0.6, 5, 14.46875],
[4, 16, 0.5, 7, 14.78125]]
for xl, xu, e_tolerance, max_iterations, Answer in QandAs: 
    result=bisection(func, xl, xu, e_tolerance, max_iterations)
    studentAnswer = result.iloc[-1]['xr']
    assert np.isclose(studentAnswer, Answer), f' Did not work for the inputs {xl=},{xu=}, {e_tolerance=}, {max_iterations=} . Expected {Answer=}. Got {studentAnswer=}'
print("All correct. Good work!")
```

## False position method
Write a function called **falsePosition** that takes the following inputs
* **f** the equation whose root is to be determined
* **xl** the lower bound of the search interval
* **xu** the upper bound of the search interval
* **eTolerance** the error tolerance in percentages
* **maxIteration** the maximum numbers of iterations

It then returns a Pandas DataFrame called **result** which contains the results from all iterations and has the following columns.
* **iteration** the iteration number
* **xl** the lower bound of the search interval
* **xu** the upper bound of the search interval
* **xr** the new root
* **relative error** the calculated relative error

Your function needs to compute the root iteratively using the False position method until the error tolerance condition has been met or the number of iterations has exceeded the maximum number of iterations. 

Using the False position method, the new root, $x_r$ is computed as $x_r = x_u - f(x_u)\frac{(x_l-x_u)}{f(x_l)-f(x_u)}$

<!-- grade_id: falsePositon, points: 0, grade: false, solution: true -->
```python
def falsePosition(func, xl, xu, e_tolerance, max_iterations): 
    """
    Perform the false position (or regula falsi) method to find the root of a function within a given interval.

    The false position method iteratively refines the interval [xl, xu] where the root of the function is located
    by using a linear interpolation to estimate the root. The process continues until the relative error falls below
    a specified tolerance or the maximum number of iterations is reached.

    Parameters:
    -----------
    func : function
        The function for which the root is to be found.
    xl : float
        The lower bound of the interval.
    xu : float
        The upper bound of the interval.
    e_tolerance : float
        The stopping criterion for the relative error (as a percentage).
    max_iterations : int
        The maximum number of iterations to perform.

    Returns:
    --------
    pandas.DataFrame
        A DataFrame containing the details of each iteration, including:
        - 'xl': The current lower bound of the interval.
        - 'xu': The current upper bound of the interval.
        - 'xr': The estimated root from the linear interpolation.
        - 'relative error': The relative error at the current iteration (as a percentage).

    Example:
    --------
    >>> def func(x):
    >>>     return x**2 - 4

    >>> falsePosition(func, 1, 3, 0.01, 100)
       xl        xu   xr        relative error
    1  1.0       3.0  1.666667  79.999999
    2  1.666667  3.0  1.846154  9.723958
    ...

    Notes:
    ------
    - The function assumes that func(xl) and func(xu) have opposite signs, indicating that a root lies within the interval [xl, xu].
    - The loop will terminate early if the function converges before reaching the maximum number of iterations.
    - The function uses the `relativeError` function to calculate the relative error.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4c2234c777536fd4, points: 0, grade: true, solution: true -->
```python
# You can call and test your function here
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: falsePositon_test, points: 10, grade: true, solution: false -->
```python
func = lambda x: 667.38/x*(1-np.exp(-0.146843*x))-40;
QandAs = [[3, 17, 0.4, 6, 14.808115949203586],
[9, 22, 0.8, 8, 14.796817165557389],
[7, 24, 0.7, 9, 14.798990111607184],
[5, 20, 0.8, 8, 14.809667776259092],
[3, 22, 0.5, 5, 14.874842878566048]]
for xl, xu, e_tolerance, max_iterations, Answer in QandAs: 
    result=falsePosition(func, xl, xu, e_tolerance, max_iterations)
    studentAnswer = result.iloc[-1]['xr']
    assert np.isclose(studentAnswer, Answer), f' Did not work for the inputs {xl=},{xu=}, {e_tolerance=}, {max_iterations=} . Expected {Answer=}. Got {studentAnswer=}'
print("All correct. Good work!")
```

## Open methods


## Newton Raphson method
Write a function called **newtonRaphson** that takes the following inputs
* **f** the equation whose root is to be determined
* **df** the derivative of the equation whose root is to be determined
* **x0** the initial guess
* **eTolerance** the error tolerance in percentages
* **maxIteration** the maximum numbers of iterations

It then returns a Pandas DataFrame called **result** which contains the results from all iterations and has the following columns.
* **iteration** the iteration number
* **x0** the old value
* **xr** the new root
* **relative error** the calculated relative error

Your function needs to compute the root iteratively using the Newton Raphson method until the error tolerance condition has been met or the number of iterations has exceeded the maximum number of iterations. 


Starting with an initial guess of $x_i = x_0$

Root is computed as $x_{i+1} = x_i - \frac{f(x_i)}{f'(x_i)}$

<!-- grade_id: newton, points: 0, grade: false, solution: true -->
```python
# Newton Raphson Method
def newtonRaphson(f, df, x0, e_tolerance, max_iterations):
    """
    Perform the Newton-Raphson method to find the root of a function.

    The Newton-Raphson method is an iterative technique that approximates the root of a function by using the function's
    derivative. The method updates the current estimate using the formula:
    
        xr = xr - f(xr) / f'(xr)
    
    The iteration continues until the relative error falls below a specified tolerance or the maximum number of iterations is reached.

    Parameters:
    -----------
    f : function
        The function for which the root is to be found.
    df : function
        The derivative of the function `f`.
    x0 : float
        The initial guess for the root.
    e_tolerance : float
        The stopping criterion for the relative error (as a percentage).
    max_iterations : int
        The maximum number of iterations to perform.

    Returns:
    --------
    pandas.DataFrame
        A DataFrame containing the details of each iteration, including:
        - 'iteration': The current iteration number.
        - 'x0': The initial guess or the previous value of xr.
        - 'xr': The updated estimate of the root.
        - 'Relative error': The relative error at the current iteration (as a percentage).

    Example:
    --------
    >>> def func(x):
    >>>     return x**2 - 4
    
    >>> def dfunc(x):
    >>>     return 2*x

    >>> newtonRaphson(func, dfunc, 1.0, 0.01, 10)
       iteration   x0   xr  Relative error
    1           1  1.0  2.0        50.000000
    2           2  2.0  1.5        33.333333
    ...

    Notes:
    ------
    - The function assumes that the derivative `df` is non-zero at the initial guess `x0`.
    - The loop will terminate early if the function converges before reaching the maximum number of iterations.
    - The function uses the `relativeError` function to calculate the relative error.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-20121573d7b66a80, points: 0, grade: true, solution: true -->
```python
# You can call and test your function here
# YOUR CODE HERE
raise NotImplementedError()
```

### Tests

<!-- grade_id: newton_test, points: 10, grade: true, solution: false -->
```python
func =  lambda x: x**3 - 6*x**2 + 11*x - 6.1
dfunc = lambda x: 3*x**2 - 12*x + 11

QandAs = [[7.5, 0.5, 5, 3.120906254949448],
[3.5, 0.5, 4, 3.0466810868815104],
[4.5, 0.7, 5, 3.0467252337803936],
[4.5, 0.3, 9, 3.0467252337803936],
[7.5, 0.4, 4, 3.342808170433576],]
for x0, e_tolerance, max_iterations, Answer in QandAs: 
    result = newtonRaphson(func, dfunc, x0, e_tolerance, max_iterations)
    studentAnswer = result.iloc[-1]['xr']
    assert np.isclose(studentAnswer, Answer), f' Did not work for the inputs {xl=},{xu=}, {e_tolerance=}, {max_iterations=} . Expected {Answer=}. Got {studentAnswer=}'
print("All correct. Good work!")
```

## Secant method

Write a function called **secant** that takes the following inputs
* **f** the equation whose root is to be determined
* **xMinus1** the first initial guess
* **x0** the second initial guess
* **eTolerance** the error tolerance in percentages
* **maxIteration** the maximum numbers of iterations

It then returns a Pandas DataFrame called **result** which contains the results from all iterations and has the following columns.
* **iteration** the iteration number
* **xMinus1** the first guess
* **x0** the second guess
* **xr** the new root
* **relative error** the calculated relative error

Your function needs to compute the root iteratively using the Secant method until the error tolerance condition has been met or the number of iterations has exceeded the maximum number of iterations. 

Root is computed as $x_{i+1} = x_i - f(x_i)\frac{x_{i-1}-x_i}{f(x_{i-1})-f(x_i)}$

### Function

<!-- grade_id: secant, points: 0, grade: false, solution: true -->
```python
def Secant(func, x0, xminus1, e_tolerance, max_iterations):
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3d834ddfcc6501f1, points: 0, grade: true, solution: true -->
```python
# You can call and test your function here
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: secant_test, points: 10, grade: true, solution: false -->
```python
func =  lambda x: x**3 - 6*x**2 + 11*x - 6.1

QandAs = [[4.5, 2.2, 0.7, 9, 1.8989430484219065],
[4.5, 1.8, 0.2, 5, 1.8989662326487895],
[5.5, 1.2, 0.6, 5, 1.0546918087104296],
[3.5, 2.0, 0.8, 7, 1.8989572776922448],
[5.5, 2.1, 0.4, 4, 1.8988210464514217],]
for x0, xminus1, e_tolerance, max_iterations, Answer in QandAs: 
    result = Secant(func, x0, xminus1, e_tolerance, max_iterations)
    studentAnswer = result.iloc[-1]['xr']
    assert np.isclose(studentAnswer, Answer), f' Did not work for the inputs {xl=},{xu=}, {e_tolerance=}, {max_iterations=} . Expected {Answer=}. Got {studentAnswer=}'
print("All correct. Good work!")
```

# Solve textbook problems
- Solve problem 5.4 using only the Bisection function you developed.
- Solve problem 6.9 using only the Newton Raphson function you developed.
- Write your scripts in the place provided below.

For all textbook problems, report your 
1. final root, 
2. the value of the function with the root, 
3. the relative error achieved, 
4. the max iterations needed, and 
5. a short conclusion about what your learnt.


<!-- grade_id: cell-2afdc0a9bb6ff969, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8e203068ee064720, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```