# MEEN 357 - Summer 2025
### Submission Instructions

- **Run All Cells**: Before submitting, go to **Kernel > Restart Kernel & Run All Cells** to ensure your code runs without errors. Submissions with errors will receive a **ZERO** grade.
- **Enter Your Name**: Fill in your name in the provided cell.
- **Complete the Code**: Replace all instances of `YOUR CODE HERE` with your solution. Remove `raise NotImplementedError()`.
- **Maintain Structure**: Do not add or remove any cells.
- **Test Your Code**: Run the provided tests to check your answers. Note that additional hidden tests may be used during grading.
- **Partial Credit**: Will be awarded only if your code runs error-free.
- **Save and Submit**: Ensure you submit the latest, correct version of your assignment by checking the last modified time.


```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3.8, "Your version of IPython is too old, please update it."
```

---

# Quiz 7 - Programming
* [Backward Substitution(6 pts)](#Backward-Substitution)
* [Multiple application of Simpson's 3/8 rule(8 pts)](#Multiple-application-of-Simpson's-3/8-rule)
* [Numerical differentiation(6 pts)](#Numerical-differentiation)


```python
import numpy as np
import pandas as pd
```

<span style='color:red '>
    
## Backward Substitution
    
</span>

Write a function called **backwardSubstitution(A, b)** that takes the following input:
* **A**, the  upper triangle found after the forward elimination method. 
$ A = \left[\matrix{a_{11}& a_{12}& a_{13} \\
         0& a'_{22}& a'_{23} \\
         0& 0& a''_{33}} \right]$
* **b**, the right hand side vector after the forward elimination method.
$ b = \left[\matrix{b_{1} \\ b'_{2} \\ b''_{3}} \right]$

and then outputs:
* **x**, the solution vector found using backward substitution.

The steps for backward substition for a 3 equations problem is shown below

$ x_3 = \frac{b_3}{a_{33}}$

$ x_2 = \frac{b_2 - a_{23}x_3}{a_{22}}$

$ x_1 = \frac{b_1 - a_{12}x_2 - a_{13}x_3}{a_{11}}$

<!-- grade_id: cell-118910e3ee38d116, points: 0, grade: false, solution: true -->
```python
def backwardSubstitution(A,b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<span style='color:red '>
    
### Script
    
</span>

**OPTIONAL:** Test your function by writing a script below.


<!-- grade_id: cell-d827e47eaf9d18f2, points: 0, grade: true, solution: true -->
```python
# Test your code here
A  = np.array([[3, -0.1, -0.2], [0, 7, -0.3], [0, 0, 10]])
b = np.array([  7.85      , -19.56166667,  70.08429319])
# YOUR CODE HERE
raise NotImplementedError()
```

<span style='color:red '>
    
### Tests
    
</span>

<!-- grade_id: cell-8eadbc13fd91d3e5, points: 6, grade: true, solution: false -->
```python
A  = np.array([[3, -0.1, -0.2], [0, 7, -0.3], [0, 0, 10]])
QandAs = [[np.array([ 14.915, -36.67 , 135.66 ]), np.array([ 5.72082762, -4.65717143, 13.566     ])],
        [ np.array([  6.28, -15.44,  57.12]), np.array([ 2.40876952, -1.96091429,  5.712     ])],
        [ np.array([  55.735, -137.03 ,  506.94 ]), np.array([ 21.37782952, -17.40311429,  50.694     ])],
        [ np.array([ 16.485, -40.53 , 149.94 ]), np.array([ 6.32302, -5.1474 , 14.994  ])],
        [ np.array([  43.175, -106.15 ,  392.7  ]), np.array([ 16.56029048, -13.48128571,  39.27      ])]]
for b, Answer in QandAs: 
    studentAnswer = backwardSubstitution(A.copy(),b)
    studentAnswer = np.ravel(studentAnswer)
    assert np.allclose(studentAnswer, Answer), f' Did not work for the inputs \n{b=}. \n\nExpected {Answer=}. \nGot {studentAnswer=}'
print("All correct. Good work")
```


<span style='color:red '>
    
## Multiple application of Simpson's 3/8 rule
    
</span>

Write a function called **simpson38** that takes the following inputs:
* **f**, the function to integrate.
* **a**, the lower limit of integration.
* **b**, the upper limit of integration.
* **n**, the number of segments.

And then outputs the following:
* **I**, the integration calculated by applying the Simpson's 3/8th rule on n segments of the function, f from a to b.

The formula for the Simpson's 3/8th Rule is:

$I = (b-a)\frac{f(x_0)+3f(x_1)+3f(x_2)+f(x_3)}{8}$

<!-- grade_id: stress, points: 0, grade: false, solution: true -->
```python
def simpson38(f, a, b, n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<span style='color:red '>
    
#### Script
    
</span>

**OPTIONAL:** Test your function by writing a script below.

<!-- grade_id: cell-754c6e5e9dddfd7a, points: 0, grade: true, solution: true -->
```python
f = np.poly1d([400, -900, 675, -200, 25, 0.2])
a = 0
b= 1
n = 3
# YOUR CODE HERE
raise NotImplementedError()
```

<span style='color:red '>
    
### Tests
    
</span>

<!-- grade_id: test_stress, points: 8, grade: true, solution: false -->
```python
f = np.poly1d([400, -900, 675, -200, 25, 0.2])
n=3
QandAs = [[0.99,1.7,170.34161446080836],
[0.52,1.1,0.9701157770600877],
[0.75,1.14,0.8725235401999979],
[0.06,1.78,262.655681030227],
[0.02,1.7,172.27049307875524]]
for a, b, Answer in QandAs: 
    studentAnswer = simpson38(f,a,b,n)
    assert np.isclose(studentAnswer, Answer), f' Did not work for the inputs \n{a=},\n{b=}. \n\nExpected {Answer=}. \nGot {studentAnswer=}'
print("All correct. Good work!")
```

<span style='color:red '>
    
## Numerical differentiation 
    
</span>


Write a function called **finiteDifference_low** that takes the following input:
* **f**, the function whose derivative we need to find,
* **x**, the point to evaluate the derivative,
* **h**, the step size, 
* **method**, the method with which to evaluate. It can be any of the following: `'forward'`, `'backward'`, or `'centered'`

And returns the following output:

* **D** the first derivative calculated with the **low accuracy formula** based on the user's choice of method. 

For example:

`D = finiteDifference_low(f, x, h, 'centered')`

Would return the first derivative of the function f, at x with a step size of h, computed using the low accuracy Centered difference formula.

The necessary formulas are shown here:

* Forward: $ f'(x_i) = \frac{f(x_{i+1})-f(x_{i})}{h} $
* Backward: $ f'(x_i) = \frac{f(x_{i})-f(x_{i-1})}{h} $
* Centered: $ f'(x_i) = \frac{f(x_{i+1})-f(x_{i-1})}{2h} $

<!-- grade_id: cell-c269d816cb6c6e47, points: 0, grade: false, solution: true -->
```python
# Numerical differentiation
def finiteDifference_low(f, x, h, method):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<span style='color:red '>
    
#### Script
    
</span>

**OPTIONAL:** Test your function by writing a script below.

<!-- grade_id: cell-d234dcc6aa8c60aa, points: 0, grade: true, solution: true -->
```python
# Test your script here
# YOUR CODE HERE
raise NotImplementedError()
```

<span style='color:red '>
    
### Tests
    
</span>

<!-- grade_id: cell-efbde15182bd40c9, points: 6, grade: true, solution: false -->
```python
f = np.poly1d([-0.1, -0.15, -0.5, -0.25, 1.2])
QandAs = [[0.44, 0.031, 'centered', -0.8115068859999987],
[0.27, 0.034, 'forward', -0.5835985383999999],
[0.28, 0.073, 'forward', -0.6246278196999998],
[0.4, 0.089, 'backward', -0.6809210131000022],
[0.48, 0.068, 'forward', -0.9376179712000011]]
for x, h, m, Answer in QandAs: 
    studentAnswer = finiteDifference_low(f, x, h, m)
    assert np.isclose(studentAnswer, Answer), f' Did not work for the inputs {x=}\n,{h=}\n,{m=}.\n\nExpected {Answer=}. \nGot {studentAnswer=}'
print("All correct. Good work!")
```