# MEEN 357 - Summer 2025
### Submission Instructions

- **Run All Cells**: Before submitting, go to **Kernel > Restart Kernel & Run All Cells** to ensure your code runs without errors. Submissions with errors will receive a **ZERO** grade.
- **Enter Your Name**: Fill in your name in the provided cell.
- **Complete the Code**: Replace all instances of `YOUR CODE HERE` with your solution. Remove `raise NotImplementedError()`.
- **Maintain Structure**: Do not add or remove any cells.
- **Test Your Code**: Run the provided tests to check your answers. Note that additional hidden tests may be used during grading.
- **Partial Credit**: Will be awarded only if your code runs error-free.
- **Save and Submit**: Ensure you submit the latest, correct version of your assignment by checking the last modified time.


```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3.8, "Your version of IPython is too old, please update it."
```

---

# Numerical Integration and Differentiation


### Trapezoidal integration
* [Multiple application of Trapezoidal Rule](#Multiple-application-of-Trapezoidal-Rule) (10 points)

### Simpson's rule
* [Multiple application of Simpson's 1/3 rule](#Multiple-application-of-Simpson's-1/3-rule) (10 points)
* [Multiple application of Simpson's 3/8 rule](#Multiple-application-of-Simpson's-3/8-rule) (10 points)

### Compare with true solution
* [True relative error in the methods](#True-relative-error-in-the-methods) (10 points)

### Finite Difference 
* [Numerical Differentiation](#Numerical-differentiation) (10 points)

### Compare with true solution
* [True solution](#Compare-with-the-true-Solution) (10 points)


```python
import numpy as np
import sympy as sp
import pandas as pd
```

## Multiple Application of the Trapezoidal Rule

Write a function called **trapezoidal** that takes the following inputs:
- **f**: The function to integrate.
- **a**: The lower limit of integration.
- **b**: The upper limit of integration.
- **n**: The number of segments (subintervals) to divide the range \([a, b]\).

The function should return:
- **I**: The estimated value of the definite integral of the function **f** from **a** to **b**, computed using the trapezoidal rule with **n** segments.

### Trapezoidal Rule Formula:
The trapezoidal rule approximates the area under the curve by dividing the range \([a, b]\) into **n** equal-width segments and calculating the area of the trapezoid for each segment. The more segments, the more accurate the approximation.

For a single segment between points **a** and **b**, the trapezoidal rule is given by:

$$
I = (b - a) \frac{f(a) + f(b)}{2}
$$

For **n** segments, the function must apply this formula repeatedly over the range \([a, b]\), summing the areas of the individual trapezoids.

### Requirements:
- The function should work for any continuous function **f**.
- Use a loop to divide the interval \([a, b]\) into **n** equal-width subintervals and apply the trapezoidal rule for each segment.
- Ensure that the final result is the sum of all these subintervals.

### Example:
```python
f = lambda x: x**2
a = 0
b = 1
n = 10
I = trapezoidal(f, a, b, n)
print(I)  # Should print the approximate value of the integral of f from 0 to 1.
```


<!-- grade_id: trapezoidalSingle, points: 0, grade: false, solution: true -->
```python
def trapezoidal(f, a, b, n):
    """
    Approximates the integral of a given function f over the interval [a, b] using the trapezoidal rule.

    Parameters:
    -----------
    f : function
        A Python function representing the mathematical function to be integrated.
    a : float
        The start of the interval over which to integrate.
    b : float
        The end of the interval over which to integrate.
    n : int
        The number of subintervals to use for the approximation.

    Returns:
    --------
    float
        The approximate value of the integral of f over [a, b] using the trapezoidal rule.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: trapezoidalSingle_script, points: 0, grade: false, solution: true -->
```python
# You can call and test your function here
f = lambda x : 400*x**5 - 900*x**4 + 675*x**3 - 200*x**2 + 25*x - 0.2
a = 0
b= 1
n = 10
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: trapezoidalSingle_test, points: 10, grade: true, solution: false -->
```python
def test_trapezoidal():
    # Define the function and test cases
    f = np.poly1d([400, -900, 675, -200, 25, 0.2])
    n = 10
    test_cases = [
        (0.81, 1.72, 192.99537845962806),
        (0.37, 1.36, 17.68401979347861),
        (0.21, 1.2, 4.2058553603051685),
        (0.44, 1.62, 111.5788682920074),
        (0.66, 1.94, 569.4407765557228)
    ]
    
    # Loop through test cases and check each one
    for i, (a, b, expected) in enumerate(test_cases, 1):
        student_answer = trapezoidal(f, a, b, n)
        assert np.isclose(student_answer, expected), (
            f'Test case {i} failed: For inputs a={a}, b={b}, '
            f'expected {expected}, but got {student_answer}.'
        )
    
    print("All tests passed successfully!")

# Call the test function
test_trapezoidal()
```

## Multiple Application of Simpson's 1/3 Rule

Write a function called **simpson13** that takes the following inputs:
- **f**: The function to integrate.
- **a**: The lower limit of integration.
- **b**: The upper limit of integration.
- **n**: The number of segments.

The function should return:
- **I**: The estimated value of the definite integral of the function **f** from **a** to **b**, computed by applying Simpson's 1/3 rule on **n** segments.

### Approach:
Simpson’s 1/3 rule approximates the integral by fitting a parabola to each segment of the function and calculating the area under the parabola. In this method, you will apply Simpson's 1/3 rule to **each segment** individually, allowing **n** to be any positive integer.

For each segment $[x_i, x_{i+1}]$, where $x_m$ is the midpoint, the Simpson's 1/3 rule is applied as:

$$
I = \frac{x_{i+1} - x_i}{6} \left( f(x_i) + 4f(x_m) + f(x_{i+1}) \right)
$$

This formula is used repeatedly for each segment between **a** and **b**, and the results are summed to get the total approximation of the integral.

### Requirements:
- The interval $[a, b]$ should be divided into **n** segments of equal width.
- For each segment, apply Simpson’s 1/3 rule and sum the results.
- **n** can be any positive integer, not necessarily even, as you are applying Simpson's rule to each segment independently.

### Example:
```python
f = lambda x: x**2
a = 0
b = 1
n = 10  # n can be any positive integer
I = simpson13(f, a, b, n)
print(I)  # Should print the approximate value of the integral of f from 0 to 1.
```

### Notes:
- **n** can be any positive integer since Simpson’s 1/3 rule is applied on each segment individually.
- Applying Simpson's 1/3 rule on each segment provides flexibility and simplifies the implementation.
- Increasing **n** improves accuracy, especially for smooth functions.

<!-- grade_id: simpson13, points: 0, grade: false, solution: true -->
```python
def simpson13(f, a, b, n):
    """
    Approximates the definite integral of a given function f over the interval [a, b] using 
    Simpson's 1/3 rule, applied on each individual segment.

    Parameters:
    -----------
    f : function
        A Python function representing the mathematical function to be integrated.
    a : float
        The start of the interval over which to integrate.
    b : float
        The end of the interval over which to integrate.
    n : int
        The number of segments into which the interval [a, b] is divided.

    Returns:
    --------
    float
        The approximate value of the integral of f over [a, b] using Simpson's 1/3 rule 
        applied on each of the n segments.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: simpson13_script, points: 0, grade: true, solution: true -->
```python
# You can call and test your function here
f = lambda x : 400*x**5 - 900*x**4 + 675*x**3 - 200*x**2 + 25*x - 0.2
a = 0
b= 1
n = 10
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: simpson13_test, points: 10, grade: true, solution: false -->
```python
def test_simpson13():
    # Define the function and test cases
    f = np.poly1d([400, -900, 675, -200, 25, 0.2])
    n = 3
    test_cases = [
        (0.63, 1.62, 107.42727556687507),
        (0.26, 1.88, 426.8775571056001),
        (0.11, 1.78, 263.2698371968452),
        (0.91, 1.14, 0.9247925416486618),
        (0.34, 1.62, 108.54189631294169)
    ]
    
    # Loop through test cases and check each one
    for i, (a, b, expected) in enumerate(test_cases, 1):
        student_answer = simpson13(f, a, b, n)
        assert np.isclose(student_answer, expected), (
            f'Test case {i} failed: For inputs a={a}, b={b}, '
            f'expected {expected}, but got {student_answer}.'
        )
    
    print("All tests passed successfully!")

# Call the test function
test_simpson13()

```

## Multiple Application of Simpson's 3/8 Rule

Write a function called **simpson38** that takes the following inputs:
- **f**: The function to integrate.
- **a**: The lower limit of integration.
- **b**: The upper limit of integration.
- **n**: The number of segments (subintervals) to divide the range $[a, b]$.

The function should return:
- **I**: The estimated value of the definite integral of the function **f** from **a** to **b**, computed using Simpson's 3/8 rule with **n** segments.

### Approach:
Simpson’s 3/8 rule is an extension of Simpson's method that approximates the integral by fitting a cubic polynomial to each segment of the function and calculating the area under the curve. In this method, you will apply Simpson's 3/8 rule on **each segment** individually, allowing **n** to be any positive integer.

For a single application of Simpson's 3/8 rule over the interval $[x_0, x_3]$, with points $x_1$ and $x_2$ dividing the interval, the formula is:

$$
I = \frac{b-a}{8} \left( f(x_0) + 3f(x_1) + 3f(x_2) + f(x_3) \right)
$$

This formula is used repeatedly for each segment between **a** and **b**, and the results are summed to get the total approximation of the integral.

### Requirements:
- Divide the interval $[a, b]$ into **n** equal segments.
- For each segment, apply Simpson’s 3/8 rule and sum the results.
- **n** can be any positive integer, as the rule is applied to each segment individually.

### Example:
```python
f = lambda x: x**3
a = 0
b = 1
n = 3  # n can be any positive integer
I = simpson38(f, a, b, n)
print(I)  # Should print the approximate value of the integral of f from 0 to 1.
```

### Notes:
- **n** can be any positive integer since Simpson’s 3/8 rule is applied to each segment independently.
- Increasing **n** improves the accuracy of the result, especially for smooth functions.

<!-- grade_id: simpson38, points: 0, grade: false, solution: true -->
```python
def simpson38(f, a, b, n):
    """
    Approximates the definite integral of a given function f over the interval [a, b] using 
    Simpson's 3/8 rule, applied on each individual segment.

    Parameters:
    -----------
    f : function
        A Python function representing the mathematical function to be integrated.
    a : float
        The start of the interval over which to integrate.
    b : float
        The end of the interval over which to integrate.
    n : int
        The number of segments into which the interval [a, b] is divided.

    Returns:
    --------
    float
        The approximate value of the integral of f over [a, b] using Simpson's 3/8 rule 
        applied on each of the n segments.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: simpson38_script, points: 0, grade: false, solution: true -->
```python
f = np.poly1d([400, -900, 675, -200, 25, 0.2])
a = 0
b= 1
n = 3
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: simpson38_test, points: 10, grade: true, solution: false -->
```python
def test_simpson38():
    # Define the function and test cases
    f = np.poly1d([400, -900, 675, -200, 25, 0.2])
    n = 3
    test_cases = [
        (0.99, 1.7, 170.34161446080836),
        (0.52, 1.1, 0.9701157770600877),
        (0.75, 1.14, 0.8725235401999979),
        (0.06, 1.78, 262.655681030227),
        (0.02, 1.7, 172.27049307875524)
    ]
    
    # Loop through test cases and check each one
    for i, (a, b, expected) in enumerate(test_cases, 1):
        student_answer = simpson38(f, a, b, n)
        assert np.isclose(student_answer, expected), (
            f"Test case {i} failed: For inputs a={a}, b={b}, "
            f"expected {expected}, but got {student_answer}."
        )
    
    print("All tests passed successfully!")

# Call the test function
test_simpson38()

```

## True error
Write a function called **trueError** that takes the following inputs:

* **x**, the new value
* **xtrue**, the previous value

And returns the following output:

* **error**, the calculated relative error

The formula for the relative error is shown below:

$e_{true} = |\frac{x-x_{true}}{x_{true}}|100 \%$

<!-- grade_id: cell-4deb474a40e835b3, points: 0, grade: false, solution: true -->
```python
def trueError(x,xtrue):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-04a341a78a1dde05, points: 0, grade: true, solution: true -->
```python
# You can call and test your function here
# YOUR CODE HERE
raise NotImplementedError()
```

## True Relative Error in Numerical Integration Methods

### Objective:
Write a script to numerically integrate the following function from **0** to **1**:

$$
f(x) = x^{0.1}(1.2 - x)(1 - e^{20(x - 1)}) \, dx
$$

### Instructions:
1. **Integration Methods**: Use the following numerical integration methods:
   - Trapezoidal Rule
   - Simpson's 1/3 Rule
   - Simpson's 3/8 Rule

2. **Parameters**:
   - Set the number of segments (n) to **4** for all integration methods.
   - Round the results of your integration to **3 decimal places**.

3. **Variable Naming**: Store the results of each integration method in the following variables:
   - `I_trap` for the Trapezoidal Rule result
   - `I_s13` for the result using Simpson's 1/3 Rule
   - `I_s38` for the result using Simpson's 3/8 Rule

4. **True Value**: The true value of the integral for comparison is **0.602298**.

5. **Calculate True Relative Error**: After computing the integrals, calculate the true relative error for each method using the formula:

$$
e_{\text{true}} = \left| \frac{I - \text{true value}}{\text{true value}} \right| \times 100\%
$$

Where $I$ is the result of each integration method.

### Example Code Structure:
You may find the following structure helpful for organizing your script:

```python
# Define the function f(x)
def f(x):
    return x**0.1 ...

# Implement the numerical integration methods here
# Calculate the integral using the Trapezoidal rule
I_trap = ...

# Calculate the integral using Simpson's 1/3 Rule
I_s13 = ...

# Calculate the integral using Simpson's 3/8 Rule
I_s38 = ...

# Print the results
print(f"Trapezoidal Rule Result: {I_trap}")
print(f"Simpson's 1/3 Rule Result: {I_s13}")
print(f"Simpson's 3/8 Rule Result: {I_s38}")

# Calculate and print the true relative error for each method
error_trap = ...
error_s13 = ...
error_s38 = ...
```

<!-- grade_id: trueCompare, points: 0, grade: false, solution: true -->
```python
# Write your script here.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b07ad544d20b8d97, points: 7, grade: true, solution: false -->
```python
# Check integration
def test_integration_results():
    # Define the expected values (rounded to 3 decimals for comparison)
    expected_values = {
        'I_trap': 0.479,  # Rounded to 3 decimal places
        'I_s13': 0.570,   # Rounded to 3 decimal places
        'I_s38': 0.579    # Rounded to 3 decimal places
    }

    # Dictionary to hold actual results
    actual_values = {
        'I_trap': I_trap,
        'I_s13': I_s13,
        'I_s38': I_s38
    }

    # Loop through each method and assert closeness to expected values
    for method, expected in expected_values.items():
        actual = actual_values[method]
        assert np.isclose(actual, expected, atol=0.001), (
            f"Test failed for {method}: Expected {expected}, but got {actual}."
        )

    print("All tests passed successfully!")

# Call the test function
test_integration_results()

```

<!-- grade_id: cell-ba6f64877644782a, points: 3, grade: true, solution: false -->
```python
# Check errors
def test_integration_results():
    # Define the expected values (rounded to 3 decimals for comparison)
    expected_values = {
        'error_trap': 20.537,  # Rounded to 3 decimal places
        'error_s13': 5.336,   # Rounded to 3 decimal places
        'error_s38': 3.87    # Rounded to 3 decimal places
    }

    # Dictionary to hold actual results
    actual_values = {
        'error_trap': error_trap,
        'error_s13': error_s13,
        'error_s38': error_s38
    }

    # Loop through each method and assert closeness to expected values
    for method, expected in expected_values.items():
        actual = actual_values[method]
        assert np.isclose(actual, expected, atol=0.001), (
            f"Test failed for {method}: Expected {expected}, but got {actual}."
        )

    print("All tests passed successfully!")

# Call the test function
test_integration_results()

```

## Numerical Differentiation

### Objective:
Write a function called **finiteDifference** to numerically approximate the first derivative of a given function at a specified point using various methods and levels of accuracy.

### Inputs:
The function should take the following parameters:
- **f**: The function for which you want to calculate the derivative. This should be a callable function (e.g., `f(x)`).
- **x**: The point at which to evaluate the derivative (a numeric value).
- **h**: The step size (a small positive numeric value).
- **method**: A string indicating the differentiation method to use. This can be one of the following options:
  - `'forward'`: Forward difference method.
  - `'backward'`: Backward difference method.
  - `'centered'`: Centered difference method.
  
- **accuracy**: A string indicating the desired accuracy of the method. This can be one of the following options:
  - `'high-order'`: High-order accuracy method.
  - `'low-order'`: Low-order accuracy method.

### Output:
The function should return:
- **D**: The first derivative of the function **f** evaluated at point **x**, computed based on the selected method and accuracy.

### Example Usage:
```python
D = finiteDifference(f, x, h, 'centered', 'high-order')
```
This would compute the derivative of the function **f** at the point **x**, using a step size of **h**, and employing the high-order centered difference formula.

### Reference:
The necessary formulas for the various differentiation methods can be found on pages 656 and 657 of your textbook. Consult these formulas to implement the appropriate calculations in your function.


<!-- grade_id: finiteDiff, points: 0, grade: false, solution: true -->
```python
# Numerical differentiation
def finiteDifference(f, x, h, method, accuracy):
    """
    Numerically approximate the first derivative of a function at a specified point.

    This function calculates the derivative of the function `f` at the point `x` 
    using finite difference methods. The method used can be selected based on 
    the desired accuracy (high-order or low-order) and the type of finite 
    difference method (forward, backward, or centered).

    Parameters:
    -----------
    f : callable
        The function for which the derivative is to be calculated. It should 
        accept a single numeric argument and return a numeric result.

    x : float
        The point at which to evaluate the derivative.

    h : float
        The step size used in the finite difference approximation. This should 
        be a small positive number.

    method : str
        The differentiation method to use. It can be one of the following:
        - 'forward': Use the forward difference method.
        - 'backward': Use the backward difference method.
        - 'centered': Use the centered difference method.

    accuracy : str
        The desired accuracy of the method. It can be one of the following:
        - 'high-order': Use a high-order accuracy method.
        - 'low-order': Use a low-order accuracy method.

    Returns:
    --------
    D : float
        The approximated first derivative of the function `f` at the point `x`, 
        calculated based on the specified method and accuracy. If an invalid 
        method or accuracy is provided, D will be set to -1.

    Raises:
    -------
    ValueError: If the method or accuracy provided is invalid.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: finiteDiff_script, points: 0, grade: false, solution: true -->
```python
# Test your script here
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: finiteDiff_test, points: 8, grade: true, solution: false -->
```python
# Define the polynomial function
f = np.poly1d([-0.1, -0.15, -0.5, -0.25, 1.2])

# Test cases: (x, h, method, accuracy, expected_answer)
test_cases = [
    (0.34, 0.032, 'centered', 'high-order', -0.6577416000000015),
    (0.45, 0.062, 'forward', 'low-order', -0.879955352800001),
    (0.52, 0.031, 'centered', 'high-order', -0.9479232000000009),
    (0.25, 0.045, 'backward', 'high-order', -0.5334171749999974),
    (0.23, 0.08, 'centered', 'low-order', -0.5102205999999998),
]

# Iterate through test cases and validate results
for i, (x, h, method, accuracy, expected) in enumerate(test_cases, 1):
    student_answer = finiteDifference(f, x, h, method, accuracy)
    
    # Assert if the student's answer is not close to the expected answer
    assert np.isclose(student_answer, expected), (
        f"Test case {i} failed for inputs: x={x}, h={h}, method='{method}', accuracy='{accuracy}'. "
        f"Expected: {expected}, Got: {student_answer}"
    )

print("All tests passed. Good work!")

```

<!-- grade_id: cell-25cf28fe8efe1b20, points: 2, grade: true, solution: false -->
```python
# Invalid test cases for method and accuracy
invalid_methods = [
    (0.34, 0.032, 'invalid_method', 'high-order'),
    (0.25, 0.045, 'backward', 'invalid_accuracy'),
]

for x, h, method, accuracy in invalid_methods:
    try:
        finiteDifference(f, x, h, method, accuracy)
    except ValueError as e:
        print(f"Correctly caught ValueError for method, and accuracy. Error message: {e}")

```

## Comparing Numerical Derivatives with the True Solution

Your task is to find the derivative of the following function at $x = 0.5$ using a series of step sizes $h$:

$$
f(x) = -0.1 x^4 - 0.15 x^3 - 0.5 x^2 - 0.25 x + 1.2
$$

### Requirements:

1. **Symbolic Derivation**:
   - Use the `sympy` library to find the exact derivative of the function. Utilize the `sp.diff()` function for differentiation.

2. **Numerical Derivation**:
   - Calculate the derivative using the **low accuracy centered difference formula**.

3. **Calculate True Relative Error**:
   - Compute the true relative error between the numerical derivative and the exact derivative.

### Step Size Range:
Generate results for step sizes ranging from $h = 0.5$ to $h = 0.05$, decrementing by $0.05$ at each step.

### Data Storage:
Create a Pandas DataFrame called **result** with the following columns:
- **step size**: The chosen step sizes $h$
- **derivative**: The numerical derivative using the low accuracy centered difference method
- **error**: The calculated true relative error for each step size


### Expected Outcome:
By the end of this exercise, you should have a comprehensive DataFrame that captures the results of your numerical differentiation, allowing for easy comparison between the different step sizes and their associated errors.


<!-- grade_id: compareToTrue, points: 0, grade: false, solution: true -->
```python
# Write your script here
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4fcb7fadd54d06ed, points: 7, grade: true, solution: false -->
```python
# Expected values for specific step sizes
expected_values = {
    0.50: (-1.0, 9.589041),
    0.45: (-0.983, 7.767123),
    0.40: (-0.968, 6.136986),
}

# Test code
for step_size, (expected_derivative, expected_error) in expected_values.items():
    student_derivative = result.loc[step_size, 'derivative']
    
    assert np.isclose(student_derivative, expected_derivative, atol=1e-3), \
        f'Failed for step size {step_size}: expected derivative {expected_derivative}, got {student_derivative}'
print("All tests passed. Good work!")
```

<!-- grade_id: cell-d7240e30e34a0c4b, points: 3, grade: true, solution: false -->
```python
# Expected values for specific step sizes
expected_values = {
    0.50: (-1.0, 9.589041),
    0.45: (-0.983, 7.767123),
    0.40: (-0.968, 6.136986),
}

# Test code
for step_size, (expected_derivative, expected_error) in expected_values.items():
    student_error = result.loc[step_size, 'error']

    assert np.isclose(student_error, expected_error, atol=1e-3), \
        f'Failed for step size {step_size}: expected error {expected_error}, got {student_error}'

print("All tests passed. Good work!")
```