# MEEN 357 - Summer 2025
### Submission Instructions

- **Run All Cells**: Before submitting, go to **Kernel > Restart Kernel & Run All Cells** to ensure your code runs without errors. Submissions with errors will receive a **ZERO** grade.
- **Enter Your Name**: Fill in your name in the provided cell.
- **Complete the Code**: Replace all instances of `YOUR CODE HERE` with your solution. Remove `raise NotImplementedError()`.
- **Maintain Structure**: Do not add or remove any cells.
- **Test Your Code**: Run the provided tests to check your answers. Note that additional hidden tests may be used during grading.
- **Partial Credit**: Will be awarded only if your code runs error-free.
- **Save and Submit**: Ensure you submit the latest, correct version of your assignment by checking the last modified time.


```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3.8, "Your version of IPython is too old, please update it."
```

---

# Quiz 3 - Programming
* [Stress Analysis of Bracket Designs](#Stress-Analysis-of-Bracket-Designs) (4 points)
* [Bisection method](#Bisection-method) (6 points)


```python
import numpy as np
import pandas as pd
```

# Stress Analysis of Bracket Designs

You will analyze a set of steel bracket designs under **axial loading**. Each bracket is subject to a force and has a known **cross-sectional area**. You will calculate the **axial stress** and determine whether the bracket design is **SAFE** or will **YIELD**, based on the material's yield strength.

Each design is loaded axially and has a rectangular cross section. The table below shows:
- `Load_N`: Applied axial force in **Newtons (N)**
- `Area_mm2`: Cross-sectional area in **square millimeters (mm²)**


The **yield strength** of the material is **250 MPa**.

Your final result should be a pandas DataFrame called `result` that includes the original data, along with two new columns:

- `Stress_MPa`: the calculated axial stress in megapascals (MPa), based on the applied load and cross-sectional area.
- `Status`: a string label — either "SAFE" or "YIELDS" — based on whether the stress is below or above the yield strength of 250 MPa.

Example output dataframe.

| Design | Load\_N | Area\_mm2 | Stress\_MPa | Status |
| ------ | ------- | --------- | ----------- | ------ |
| E      | 18000    | 180       | 100.0       | SAFE   |
| F      | 21000    | 140       | 150.0       | SAFE   |
| G      | 25000    | 100       | 250.0       | YIELDS |
| H      | 30000    | 90        | 333.3       | YIELDS |


**HINT**: In pandas, you can perform calculations across columns without using `apply()`. See sample code below.

`df["Volume_cm3"] = df["Length_cm"] * df["Width_cm"] * df["Height_cm"]`



```python
# Given design table. Do not remove or modify.
data = {
    "Design": ["A", "B", "C", "D"],
    "Load_N": [15000, 22500, 18000, 24000],
    "Area_mm2": [150, 75, 100, 60],  # mm²
}

df = pd.DataFrame(data)
print(df)
```

<!-- grade_id: stress, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: test_stress, points: 2, grade: true, solution: false -->
```python
# Check Design C's stress value
stress_C = df.loc[df["Design"] == "C", "Stress_MPa"].values[0]
assert round(stress_C, 1) == 180.0, f"Incorrect stress for Design C: expected 180.0, got {stress_C}"
print('Stress calculation is correct')
```

<!-- grade_id: test_status, points: 2, grade: true, solution: false -->
```python
# Check Design D's status
status_D = df.loc[df["Design"] == "D", "Status"].values[0]
assert status_D == "YIELDS", f"Incorrect status for Design D: expected 'YIELDS', got '{status_D}'"
print('Status is correct')

```

## Bisection method
Write a function called “bisection” that takes the following inputs
* **f** the equation whose root is to be determined
* **xl** the lower bound of the search interval
* **xu** the upper bound of the search interval
* **max_iteration** the maximum numbers of iterations

It then returns a Pandas DataFrame called **result**, which contains the results from all iterations and has the following columns.
* **iteration** the iteration number
* **xl** the lower bound of the search interval
* **xu** the upper bound of the search interval
* **xr** the new root

Your function needs to compute the root iteratively using the bisection method until the number of iterations has exceeded the maximum number of iterations. 

Using the bisection method, the new root, $x_r$ is computed as $x_r = \frac{x_l+x_u}{2}$

<!-- grade_id: bisection, points: 0, grade: false, solution: true -->
```python
def bisection(f, xl, xu, max_iterations): 
    """
    Perform the bisection method to find the root of a function within a given interval.

    The bisection method iteratively narrows the interval [xl, xu] where the root of the function f(x) is located.
    It stops when the maximum number of iterations is reached.

    Parameters:
    -----------
    f : function
        The function for which the root is to be found.
    xl : float
        The lower bound of the interval.
    xu : float
        The upper bound of the interval.
    max_iterations : int
        The maximum number of iterations to perform.

    Returns:
    --------
    pandas.DataFrame
        A DataFrame containing the details of each iteration, including:
        - 'xl': The current lower bound of the interval.
        - 'xu': The current upper bound of the interval.
        - 'xr': The midpoint of the current interval (the approximate root).

    Example:
    --------
    >>> def func(x):
    >>>     return x**2 - 4

    >>> bisection(func, 1, 3, 100)
          xl   xu   xr  
    1     1.0  3.0  2.0
    2     1.0  2.0  1.5
    ...

    Notes:
    ------
    - The function assumes that f(xl) and f(xu) have opposite signs, indicating that a root lies within the interval [xl, xu].

    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ec3e7957b793484d, points: 0, grade: true, solution: true -->
```python
# You can call and test your function here

f = lambda x: 667.38/x*(1-np.exp(-0.146843*x))-40;
xl = 12
xu = 16
max_iterations = 10
result = bisection(f, xl, xu, max_iterations)
result

# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: test_if_error, points: 2, grade: true, solution: false -->
```python
f = lambda x: 667.38/x*(1-np.exp(-0.146843*x))-40;
xl = 12
xu = 16
max_iterations = 10
result=bisection(f, xl, xu, max_iterations)
assert all("err" not in col.lower() for col in result.columns), "Did not ask for relative error."
```

<!-- grade_id: bisection_test, points: 4, grade: true, solution: false -->
```python
f = lambda x: 667.38/x*(1-np.exp(-0.146843*x))-40;
QandAs = [[3, 18, 0.3, 6, 14.953125],
[7, 17, 0.2, 8, 14.7734375],
[10, 23, 0.4, 5, 14.46875],
[10, 21, 0.6, 5, 14.46875],
[4, 16, 0.5, 7, 14.78125]]
for xl, xu, e_tolerance, max_iterations, Answer in QandAs: 
    result=bisection(f, xl, xu, max_iterations)
    studentAnswer = result.iloc[-1]['xr']
    assert np.isclose(studentAnswer, Answer), f' Did not work for the inputs {xl=},{xu=}, {max_iterations=} . Expected {Answer=}. Got {studentAnswer=}'
print("Root calculation is correct. Good work!")
```