# MEEN 357 - Summer 2025
### Submission Instructions

- **Run All Cells**: Before submitting, go to **Kernel > Restart Kernel & Run All Cells** to ensure your code runs without errors. Submissions with errors will receive a **ZERO** grade.
- **Enter Your Name**: Fill in your name in the provided cell.
- **Complete the Code**: Replace all instances of `YOUR CODE HERE` with your solution. Remove `raise NotImplementedError()`.
- **Maintain Structure**: Do not add or remove any cells.
- **Test Your Code**: Run the provided tests to check your answers. Note that additional hidden tests may be used during grading.
- **Partial Credit**: Will be awarded only if your code runs error-free.
- **Save and Submit**: Ensure you submit the latest, correct version of your assignment by checking the last modified time.


```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3.8, "Your version of IPython is too old, please update it."
```

---

# Solving PDEs

### Explicit methods  
* [Explicit method](#Explicit-method) (7 points)

### Implicit methods
* [Implicit method](#Implicit-method) (7 points)

### Script
* [Solve PDE with both functions](#Solve-the-heat-Equation)  (6 points)


```python
import numpy as np
import matplotlib.pyplot as plt
```

## Explicit method

Write a function " *heatEquationExplicit* " that takes the following inputs

* **x** , the input vector x. This will be a vector of equal step size, $h$.
* **t** , the time vector t. This will be a vector of equal step size, $k$.
* **boundaryConditions** , the boundary conditions $T(x=0, t)\ and\ T(x=L, t)$
* **initialConditions** , the initial conditions $T(x,t=0)$

and returns the solution to the PDE 
$$\frac{\partial^2u}{\partial x^2} -  \frac{\partial u}{\partial t} = 0$$

computed using the Explicit method.

* **T** , the solution matrix (whose number of rows = length of position vector, $x$, and number of columns = length of time, $t$ )

The equation derived with the Explicit method for finding the Temperature in the next time value is given below.

$T(x,t+k) = \lambda\ T(x-h, t)+(1-2\lambda)\ T(x,t)+\lambda\ T(x+h, t)$

The same can be expressed in terms of the row (position) and column (time) indicies as

$T_{i,j} = \lambda\ T_{i-1, j-1}+(1-2\lambda)\ T_{i,j-1}+\lambda\ T_{i+1, j-1}$

<!-- grade_id: Explicit, points: 0, grade: false, solution: true -->
```python
def heatEquationExplicit(x, t, boundaryConditions, initialConditions):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: Explicit_script, points: 0, grade: true, solution: true -->
```python
# Test your function here.
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: Explicit_test, points: 7, grade: true, solution: false -->
```python
k = 0.025

QandA = [[0.25, np.array([0.707, 0.541, 0.415, 0.317, 0.243])],
[0.2, np.array([0.588, 0.447, 0.341, 0.259, 0.197])],
[0.1, np.array([0.309, 0.233, 0.176, 0.133, 0.101])],
[0.125, np.array([0.383, 0.289, 0.219, 0.166, 0.125])]]
for h, Answer in QandA:
    x = np.arange(0,1+h,h).round(3)
    t = np.arange(0,0.1+k,k).round(3)
    boundaryConditions = [0,0]
    initialConditions = np.sin(np.pi*x)
    T = heatEquationExplicit(x, t, boundaryConditions, initialConditions)
    studentAnswer = T[1].round(3)
    assert np.allclose(studentAnswer, Answer), f' Did not work for {h=}. \nExpected {Answer=}. \nGot {studentAnswer=}'

print('All correct. Good work !')
```

## Implicit method

Write a function " *heatEquationImplicit* " that takes the following inputs

* **x** , the input vector x. This will be a vector of equal step size, $h$.
* **t** , the time vector t. This will be a vector of equal step size, $k$.
* **boundaryConditions** , the boundary conditions $T(x=0, t)\ and\ T(x=L, t)$
* **initialConditions** , the initial conditions $T(x,t=0)$

and returns the solution to the PDE 
$$\frac{\partial^2u}{\partial x^2} -  \frac{\partial u}{\partial t} = 0$$

computed using the Implicit method.

* **T** , the solution matrix (whose number of rows = length of position vector, $x$, and number of columns = length of time, $t$ )

The equation derived with the Implicit method for finding the Temperature in the next time value is given below.

$ -\lambda\ T(x-h, t)+(1+2\lambda)\ T(x,t)-\lambda\ T(x+h, t) = T(x,t-k)$

The same can be expressed in terms of the row (position) and column (time) indicies as

$ -\lambda\ T_{i-1, j}+(1+2\lambda)\ T_{i,j}-\lambda\ T_{i+1, j} = T_{i,j-1}$

<!-- grade_id: Implicit, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: CNscript, points: 0, grade: true, solution: true -->
```python
# Test your function here.
# YOUR CODE HERE
raise NotImplementedError()
```

### Tests

<!-- grade_id: imp_test, points: 7, grade: true, solution: false -->
```python
k = 0.025

QandA = [[0.25, np.array([0.707, 0.573, 0.464, 0.376, 0.305])] ,
[0.2, np.array([0.588, 0.475, 0.383, 0.309, 0.25 ])] ,
[0.1, np.array([0.309, 0.248, 0.199, 0.16 , 0.129])] ,
[0.125, np.array([0.383, 0.308, 0.247, 0.199, 0.16 ])]]
for h, Answer in QandA:
    x = np.arange(0,1+h,h).round(3)
    t = np.arange(0,0.1+k,k).round(3)
    boundaryConditions = [0,0]
    initialConditions = np.sin(np.pi*x)
    T = heatEquationImplicit(x, t, boundaryConditions, initialConditions)
    studentAnswer = T[1].round(3)
    assert np.allclose(studentAnswer, Answer), f' Did not work for {h=}. \nExpected {Answer=}. \nGot {studentAnswer=}'

print('All correct. Good work !')
```

### Solve the heat Equation
Solve the folloing PDE using the Explicit method, and the Implicit method developed above. 
$$\frac{\partial^2u}{\partial x^2} -  \frac{\partial u}{\partial t} = 0$$

* Solve it
    * from $x = 0$ to $x=1$ with a step size of $h = 0.10$, and 
    * from $t = 0$ to $t=1$ with a step size of $k = 0.025$. 
    * The boundary conditions are, $T(0,t) = T(1,t) = 0$
    * The initial condition at $T(x,0) = sin(\pi x)$. 
* Store the result for
    * Explicit method in **T_Explicit**
    * Implicit method in **T_Implicit**

* Create 2 subplots side by side for the Explicit method and the Implicit method. Plot the temperature vs time graph for the various times on the same figure.
    * Save the plot as **PDE_Heat_equation.jpg**



<!-- grade_id: PDE_script, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9116ab7142e73143, points: 3, grade: true, solution: false -->
```python
Answer = [np.array([[0.    , 0.    , 0.    , 0.    , 0.    ],
       [0.309 , 0.2334, 0.1763, 0.1331, 0.1006],
       [0.5878, 0.4439, 0.3353, 0.2532, 0.1913],
       [0.809 , 0.611 , 0.4615, 0.3486, 0.2633],
       [0.9511, 0.7183, 0.5425, 0.4098, 0.3095],
       [1.    , 0.7553, 0.5705, 0.4309, 0.3254],
       [0.9511, 0.7183, 0.5425, 0.4098, 0.3095],
       [0.809 , 0.611 , 0.4615, 0.3486, 0.2633],
       [0.5878, 0.4439, 0.3353, 0.2532, 0.1913],
       [0.309 , 0.2334, 0.1763, 0.1331, 0.1006],
       [0.    , 0.    , 0.    , 0.    , 0.    ]]), np.array([[0.    , 0.    , 0.    , 0.    , 0.    ],
       [0.309 , 0.2483, 0.1995, 0.1602, 0.1287],
       [0.5878, 0.4722, 0.3794, 0.3048, 0.2449],
       [0.809 , 0.65  , 0.5222, 0.4195, 0.337 ],
       [0.9511, 0.7641, 0.6139, 0.4932, 0.3962],
       [1.    , 0.8034, 0.6454, 0.5185, 0.4166],
       [0.9511, 0.7641, 0.6139, 0.4932, 0.3962],
       [0.809 , 0.65  , 0.5222, 0.4195, 0.337 ],
       [0.5878, 0.4722, 0.3794, 0.3048, 0.2449],
       [0.309 , 0.2483, 0.1995, 0.1602, 0.1287],
       [0.0    , 0.    , 0.    , 0.    , 0.    ]])]
assert np.allclose(T_Explicit.round(4), Answer[0]), f' Did not generate the correct answer from the Explicit method'
assert np.allclose(T_Implicit.round(4), Answer[1]), f' Did not generate the correct answer from the Crank Nicholson method'
print('All correct. Good work !')

```