# MEEN 357 - Summer 2025
### Submission Instructions

- **Run All Cells**: Before submitting, go to **Kernel > Restart Kernel & Run All Cells** to ensure your code runs without errors. Submissions with errors will receive a **ZERO** grade.
- **Enter Your Name**: Fill in your name in the provided cell.
- **Complete the Code**: Replace all instances of `YOUR CODE HERE` with your solution. Remove `raise NotImplementedError()`.
- **Maintain Structure**: Do not add or remove any cells.
- **Test Your Code**: Run the provided tests to check your answers. Note that additional hidden tests may be used during grading.
- **Partial Credit**: Will be awarded only if your code runs error-free.
- **Save and Submit**: Ensure you submit the latest, correct version of your assignment by checking the last modified time.


```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3.8, "Your version of IPython is too old, please update it."
```

---

# Programming in Python

## Volume Calculator

Complete Problem 3.1 from the textbook.


<!-- grade_id: volume, points: 0, grade: false, solution: true -->
```python
import numpy as np
def volume_calculator(R, d):
    """ This function calculates the volume of water in a tank.
    
    Parameters
    R : float
        The radius of the cylinder.
    d : float
        Depth of water.

    Returns
    V : float
        The volume of water in the tank.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
# You can call and test your function here
R = 0.9
d = 1
V = volume_calculator(R,d)
print(V)
```

**Test code**

 Run the cell below. If it ran without any errors, your function is working correctly.

<!-- grade_id: volume_test, points: 4, grade: true, solution: false -->
```python
QandA = [[0.9, 1.0, 1.017876019763093] ,
[1.5, 1.25, 2.045307717180855] ,
[1.3, 3.8, 15.573921981395802] ,
[1.3, 3.0, 11.326488713742402]]

for entry in QandA:
    R, d, V = entry
    student_answer = volume_calculator(R, d)
    assert np.isclose(student_answer, V) , f'Wrong result for {R=}, {d=}. Expected {V=}, got {student_answer=}'

print('All correct. Good work.')
```

<!-- grade_id: volume_error, points: 1, grade: true, solution: false -->
```python
try:
    volume_calculator(1,4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError : Overtop")
```

## Cartesian coordinates to polar coordinates

Complete problem 3.6 from the textbook.


<!-- grade_id: cart, points: 0, grade: false, solution: true -->
```python
def cart_to_polar(x, y):
    """ This function converts cartesian coordinates to polar coordinates. 

    Parameters
    x : float
        x-coordinate
    y : float
        y-coordinate

    Returns
    r : Scalar
        Radial distance
    theta : Scalar
            Angle
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# You can call and test your function here
r, theta = cart_to_polar(-1,-1)

```

**Test code**

 Run the cell below. If it ran without any errors, your function is working correctly.

<!-- grade_id: cart_test, points: 5, grade: true, solution: false -->
```python
QandA = [[-1, -1, 1.4142135623730951, -2.356194490192345] ,
[-1, 0, 1.0, 3.141592653589793] ,
[-1, 1, 1.4142135623730951, 2.356194490192345] ,
[-1, 2, 2.23606797749979, 2.0344439357957027] ,
[0, -1, 1.0, -1.5707963267948966], #-3.141592653589793] ,
[0, 0, 0.0, 0] ,
[0, 1, 1.0, 1.5707963267948966] ,
[0, 2, 2.0, 1.5707963267948966] ,
[1, -1, 1.4142135623730951, -0.7853981633974483] ,
[1, 0, 1.0, 0.0] ,
[1, 1, 1.4142135623730951, 0.7853981633974483] ,
[1, 2, 2.23606797749979, 1.1071487177940904] ,
[2, -1, 2.23606797749979, -0.4636476090008061] ,
[2, 0, 2.0, 0.0] ,
[2, 1, 2.23606797749979, 0.4636476090008061] ,
[2, 2, 2.8284271247461903, 0.7853981633974483] ,]
for entry in QandA:
    x, y, r, theta = entry
    student_r, student_theta = cart_to_polar(x, y)
    assert np.isclose([r, theta], [student_r, student_theta]).all(), f'did not work for {x=}, {y=}. Expected {r=}, {theta=}. Got {student_r=}, {student_theta=}'

print('All correct. Good work.')
```

## Grade Calculator
Complete problem 3.8 from the textbook.


<!-- grade_id: grade, points: 0, grade: false, solution: true -->
```python
def lettergrade(score):
    """ This function returns the letter grade for a given score.

    Parameters
    score : float
            The score to be converted.

    Returns
    grade : String
            The calculated letter grade.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# You can call and test your function here
score = 98
grade = lettergrade(score)
print(f'The grade for a score of {score}/100 is {grade}')

```

**Test code**

 Run the cell below. If it ran without any errors, your function is working correctly.

<!-- grade_id: grade_test, points: 4, grade: true, solution: false -->
```python
QandA = [[10,'F'],
        [61, 'D'],
        [72, 'C'],
        [83, 'B'],
        [94, 'A']]
for entry in QandA:
    score, grade = entry
    student_grade = lettergrade(score)
    assert grade==student_grade, f'Did not work for {score=}. Expected {grade=}. Got {student_grade=}'

print('All correct. Good work.')
```

<!-- grade_id: grade_error, points: 1, grade: true, solution: false -->
```python
try:
    lettergrade(105)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")
```

## Import and Apply a Custom Function
Apply your custom letter grade function to process a dataset of student scores and generate a new file with letter grades.

Prerequisites:
- Basic understanding of Python.
- Familiarity with functions and importing modules.
- Knowledge of Pandas library basics.

Instructions:

1. Function Creation:
   - Create a file named `my_tools.py`.
   - In this file, implement your `lettergrade` function that converts a numeric score to a letter grade.
   - Ensure your function handles scores from 0 to 100 and returns appropriate letter grades (A, B, C, D, F).

2. Data Processing Script:
   - Write your script in the space below.
   - In this script, you will process a CSV file containing student scores.

3. Import Necessary Modules:
   - Import your `lettergrade` function from `my_tools.py`.
   - Import the `pandas` library.

4. Read the Data:
   - Use pandas to read a CSV file named 'scores.csv'.
   - This file contains a column called 'score' for students' numeric scores.

5. Apply Your Function:
   - Use the `apply` method of pandas to run your `lettergrade` function on the 'score' column.
   - Create a new column named 'grade' to store the resulting letter grades.

6. Save the Results:
   - Save the updated DataFrame to a new CSV file named 'grades.csv'.
   - Ensure that you don't include the index in the output file.

7. Testing and Verification:
   - Run your script and check the output 'grades.csv' file.
   - Verify that the letter grades are correctly assigned based on the numeric scores.

Submission:
- Submit your `my_tools.py` files

<!-- grade_id: cell-6a4a54d64a23afb8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f64c0233a3e2ee11, points: 5, grade: true, solution: false -->
```python
assert (data.iloc[range(8)]['grade']==['F','F','F','F','F','F','A','F']).all(), "Script did not work."

print('All correct. Good work.')
```