# MEEN 357 - Summer 2025
### Submission Instructions

- **Run All Cells**: Before submitting, go to **Kernel > Restart Kernel & Run All Cells** to ensure your code runs without errors. Submissions with errors will receive a **ZERO** grade.
- **Enter Your Name**: Fill in your name in the provided cell.
- **Complete the Code**: Replace all instances of `YOUR CODE HERE` with your solution. Remove `raise NotImplementedError()`.
- **Maintain Structure**: Do not add or remove any cells.
- **Test Your Code**: Run the provided tests to check your answers. Note that additional hidden tests may be used during grading.
- **Partial Credit**: Will be awarded only if your code runs error-free.
- **Save and Submit**: Ensure you submit the latest, correct version of your assignment by checking the last modified time.


```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3.8, "Your version of IPython is too old, please update it."
```

---

# Least Squares


### Polynomial Regression
* [Least Squares polynomial regression](#Least-Squares-polynomial-regression) (12 points)

### Compare with built-in methods
* [Compare with Numpy's polyfit() function](#True-relative-error-in-the-methods) (Script/8 points)


```python
import numpy as np
import sympy as sp
import pandas as pd
import matplotlib.pyplot as plt
```

## Least Squares polynomial regression

Write a function called **LeastSquares(x, y, n)** that takes the following inputs:
* **x**, the independent variable, x.
* **y**, the dependent variable, y.
* **n**, the order of the polynomial equation.

And then outputs the following:
* **c**, the coefficients of the polynomial fit equation such that

$ f(x) = c_nx^n + c_{n-1}x^{n-1} +....c_1x + c_0x $


<!-- grade_id: trapezoidalSingle, points: 0, grade: false, solution: true -->
```python
def LeastSquares(x, y, n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: trapezoidalSingle_script, points: 0, grade: true, solution: true -->
```python
# You can call and test your function here
data = pd.read_csv('data.csv')
data.sort_values('x', inplace=True)
data.head()
x = np.array(data['x'].values)
y = np.array(data['y'].values)
# YOUR CODE HERE
raise NotImplementedError()
```

### Test code

<!-- grade_id: trapezoidalSingle_test, points: 12, grade: true, solution: false -->
```python
data = pd.read_csv('data.csv')
data.sort_values('x', inplace=True)
data.head()
x = np.array(data['x'].values)
y = np.array(data['y'].values)

QandAs = [[0,np.array([72.73505055])],
[1,np.array([1.32243102, 7.99102098])],
[2,np.array([1.20413998e-02, 1.45453992e-01, 3.56192048e+01])],
[3,np.array([6.59975564e-05, 2.37046241e-03, 6.03121892e-01, 2.86471980e+01])],
[4,np.array([ 6.14036121e-05, -1.17656377e-02,  8.33648111e-01, -2.45781244e+01,
        3.05547622e+02])]]

for n, Answer in QandAs: 
    studentAnswer = LeastSquares(x, y, n)
    studentAnswer = studentAnswer.ravel()
    assert np.allclose(studentAnswer, Answer), f' Did not work for the inputs {x=}. Expected {Answer=}. Got {studentAnswer=}'
print("All correct. Good work!")
```

## Compare the result with Numpy polyfit function

Write a **script** to fit a second-order polynomial (n=2) to the following data

| | | | | | | |
|---|---|---|---|---|---|---|
|$x_i$|0|1|2|3|4|5|
|$y_i$|2.1|7.7|13.6|27.2|40.9|61.1|

using:

* The Least squares polynomial fit function you developed
* The Numpy polyfit function

Plot the following two subplots on a figure
* suplot 1
    * The given data points, **x, y**.
    * The polynomial equation found using the Least squares polynomial fit function you developed

* suplot 2
    * The given data points, **x, y**.
    * The polynomial equation found using the Numpy polyfit function

Save the figure as *"polynomialFit.png"*

<!-- grade_id: trueCompare, points: 8, grade: true, solution: true -->
```python
# Write your script here.
# YOUR CODE HERE
raise NotImplementedError()
```

## Estimate the spring stiffness from data
The data set called **spring_data.csv** contains the load and deflection information of 2 springs. Use linear regression to find the spring stiffnesses. 

Plot the linear models and the data together.

<!-- grade_id: spring_data_k, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```